Allosteric activation of the co-receptor BAK1 by the EFR receptor kinase initiates immune signaling

Author response:

The following is the authors’ response to the previous reviews.

eLife assessment

This manuscript reports important in vitro biochemical and in planta experiments to study the receptor activation mechanism of plant membrane receptor kinase complexes with non-catalytic intracellular kinase domains. Several lines of evidence convincingly show that one such putative pseudokinase, the immune receptor EFR achieves an active conformation following phosphorylation by a co-receptor kinase, and then in turn activates the co-receptor kinase allosterically to enable it to phosphorylate down-stream signaling components. This manuscript will be of interest to scientists focusing on cell signalling and allosteric regulation.

We wish to clarify that EFR is itself, not a pseudokinase. We could show in previous work (Bender et al., 2021; https://doi.org/10.1073/pnas.2108242118 ) that EFR has catalytic activity in vitro. This catalytic activity is, however, not required for elf18-induced immune signaling in planta.

Public Reviews:

Reviewer #1 (Public Review):

Summary

The authors use an elegant but somewhat artificial heterodimerisation approach to activate the isolated cytoplasmic domains of different receptor kinases (RKs) including the receptor kinase BRI1 and EFR. The developmental RK BRI1 is known to be activated by the co-receptor BAK1. Active BRI1 is then able to phosphorylate downstream substrates. The immune receptor EFR is also an active protein kinase also activated by the co-receptor BAK1. EFR however appears to have little or no kinase activity but seems to use an allosteric mechanism to in turn enable BAK1 to phosphorylate the substrate kinase BIK1. EFR tyrosine phosphorylation by BAK1 appears to trigger a conformational change in EFR, activating the receptor. Likewise, kinase activating mutations can cause similar conformational transitions in EFR and also in BAK1 in vitro and in planta.

We wish to clarify that we make no strong link between tyrosine phosphorylation and the conformational change leading to activation of the complex. Rather, the HDX-MS data demonstrate the structural importance of Tyr836 for the activation mechanism. At present, we do not know how phosphorylation of the residue would affect the activation process.

Strengths:

I particularly liked The HDX experiments coupled with mutational analysis (Fig. 2) and the design and testing of the kinase activating mutations (Fig. 3), as they provide novel mechanistic insights into the activation mechanisms of EFR and of BAK1. These findings are nicely extended by the large-scale identification of EFR-related RKs from different species with potentially similar activation mechanisms (Fig. 5).

Weaknesses:

In my opinion, there are currently two major issues with the present manuscript. (1) The authors have previously reported that the EFR kinase activity is dispensible for immune signaling (https://pubmed.ncbi.nlm.nih.gov/34531323/) but the wild-type EFR receptor still leads to a much better phosphorylation of the BIK1 substrate when compared to the kinase inactive D849N mutant protein (Fig. 1). (2) How the active-like conformation of EFR is in turn activating BAK1 is poorly characterized, but appears to be the main step in the activation of the receptor complex. Extending the HDX analyses to resting and Rap-activated receptor complexes could be a first step to address this question, but these HDX studies were not carried out due to technical limitations.

Overall this is an interesting study that aims to advance our understanding of the activation mechanisms of different plant receptor kinases with important functions in plant immunity.

Reviewer #2 (Public Review):

Summary:

Transmembrane signaling in plants is crucial for homeostasis. In this study, the authors set out to understand to what extent catalytic activity in the EFR tyrosine kinase is required in order to transmit a signal. This work was driven by mounting data that suggest many eukaryotic kinases do not rely on catalysis for signal transduction, relying instead on conformational switching to relay information. The crucial findings reported here involve the realisation that a kinase-inactive EFR can still activate (ie lead to downstream phosphorylation) of its partner protein BAK1. Using a convincing set of biochemical, mass spectrometric (HD-exchange) and in vivo assays, the team suggest a model in which EFR is likely phosphorylated in the canonical activation segment (where two Ser residues are present), which is sufficient to generate a conformation that can activate BAK1 through dimersation. A model is put forward involving C-helix positioning in BAK1, and the model extended to other 'non-RD' kinases in Arabidopsis kinases that likely do not require kinase activity for signaling.

We prefer not to describe EFR as a tyrosine kinase. It may be the case that EFR can function under certain conditions as a dual-specificity protein kinase, but this has never been demonstrated experimentally. We therefore describe EFR as a Ser/Thr protein kinase, since it is known that the isolated cytoplasmic domain can phosphorylate on Ser and Thr residues (Wang et al., 2014; https://doi.org/10.1016/j.jprot.2014.06.009).

Strengths:

The work uses logical and well-controlled approaches throughout, and is clear and convincing in most areas, linking data from IPs, kinase assays (including clear 32P-based biochemistry), HD-MX data (from non-phosphorylated EFR) structural biology, oxidative burst data and infectivity assays. Repetitions and statistical analysis all appear appropriate.

Overall, the work builds a convincing story and the discussion does a clear job of explaining the potential impact of these findings (and perhaps an explanation of why so many Arabidopsis kinases are 'pseudokinases', including XPS1 and XIIa6, where this is shown explicitly).

Weaknesses:

No major weaknesses are noted from reviewing the data and the paper follows a logical course built on solid foundations; the use of Tables to explain various experimental data pertinent to the reported studies is appreciated.

(1) The use of a, b,c, d in Figures 2C and 3C etc is confusing to this referee, and is now addressed in the latest version

(2) The debate about kinase v pseudokinases is well over a decade old. For non-experts, the kinase alignments/issues raised are in PMID: 23863165 and might prove useful if cited.

We have cited the suggested reference in the second paragraph of the discussion.

(3) Early on in the paper, the concept of kinases and pseudokinases related to R-spine (and extended R-spine) stability and regulation really needs to be more adequately introduced to explain what comes next; e.g. some of the key work in this area for RAF and Tyr kinases where mutual F-helix Phe amino acid changes are evaluated (conceptually similar to this study of the E-helix Tyr to Phe changes in EFR) should be cited (PMID: 17095602, 24567368 and 26925779).

As an alternative, we have amended the text in several places to focus on conformational toggling between active/inactive states rather than R-spine stability. We think that this keeps the message of our manuscript focused. We hope that the reviewer finds this acceptable.

(4) In my version, some of the experimental text is also currently in the wrong order (and no page numbers, so hard for me to state exactly where in the manuscript); However, I am certain that Figure 2C is mentioned in the text when the data are actually shown in Figure 3C for the EFR-SSAA protein.

Indeed, some references to Figure 2 in the text were incorrect. We have corrected these. References in the text to Figure 3 and the data reported therein are correct.

(5) Tyr 156 in PKA is not shown in Supplement 1, 2A as suggested in the text; for readers, it will be important to show the alignment of the Tyr residue in other kinases; this has been updated in the second version. Although it is clearly challenging to generate phosphorylated EFR (seemingly through Codon-expansion here?), it appears unlikely that a phosphorylated EFR protein, even semi-pure, couldn't have been assayed to test the idea that the phosphorylation drives/supports downstream signaling. What about a DD or EE mutation, as commonly used (perhaps over-used) in MEK-type studies?

Our aim with codon expansion was to generate recombinant protein carrying high-stoichiometry phosphorylation at sites which we have previously documented to be required for downstream signaling (Macho et al., 2014; Bender et al., 2021). We additionally demonstrated previously that a DD mutant of the activation loop sites in EFR does not fully complement the efr-1 mutant (Bender et al., 2021), suggesting that the Asp mutations are not good phospho-mimics in this context. We therefore did not generate DD or EE mutations for in vitro studies.

Impact:

The work is an important new step in the huge amount of follow-up work needed to examine how kinases and pseudokinases 'talk' to each other in (especially) the plant kingdom, where significant genetic expansions have occurred. The broader impact is that we might understand better how to manipulate signaling for the benefit of plants and mankind; as the authors suggest, their study is a natural progression both of their own work, and the kingdom-wide study of the Kannan group.

Reviewer #3 (Public Review):

The study presents strong evidence for allosteric activation of plant receptor kinases, which enhances our understanding of the non-catalytic mechanisms employed by this large family of receptors.

Plant receptor kinases (RKs) play a critical role in transducing extracellular signals. The activation of RKs involves homo- or heterodimerization of the RKs, and it is believed that mutual phosphorylation of their intracellular kinase domains initiates downstream signaling. However, this model faces a challenge in cases where the kinase domain exhibits pseudokinase characteristics. In their recent study, Mühlenbeck et al. reveal the non-catalytic activation mechanisms of the EFR-BAK1 complex in plant receptor kinase signaling. Specifically, they aimed to determine that the EFR kinase domain activates BAK1 not through its kinase activity, but rather by utilizing a "conformational toggle" mechanism to enter an active-like state, enabling allosteric trans-activation of BAK1. The study sought to elucidate the structural elements and mutations of EFR that affect this conformational switch, as well as explore the implications for immune signaling in plants. To investigate the activation mechanisms of the EFR-BAK1 complex, the research team employed a combination of mutational analysis, structural studies, and hydrogen-deuterium exchange mass spectrometry (HDX-MS) analysis. For instance, through HDX-MS analysis, Mühlenbeck et al. discovered that the EFR (Y836F) mutation impairs the accessibility of the active-like conformation. On the other hand, they identified the EFR (F761H) mutation as a potent intragenic suppressor capable of stabilizing the active-like conformation, highlighting the pivotal role of allosteric regulation in BAK1 kinase activation. The data obtained from this methodology strengthens their major conclusion. Moreover, the researchers propose that the allosteric activation mechanism may extend beyond the EFR-BAK1 complex, as it may also be partially conserved in the Arabidopsis LRR-RK XIIa kinases. This suggests a broader role for non-catalytic mechanisms in plant RK signaling.

The allosteric activation mechanism was demonstrated for receptor tyrosine kinases (RTKs) many years ago. A similar mechanism has been suggested for the activation of plant RKs, but experimental evidence for this conclusion is lacking. Data in this study represent a significant advancement in our understanding of non-catalytic mechanisms in plant RK signaling. By shedding light on the allosteric regulation of BAK1, the study provides a new paradigm for future research in this area.

Recommendations for the authors:

Reviewer #1 (Recommendations For The Authors):

The authors have considered points 1-5 raised in my initial review and the revised manuscript contains a more balanced discussion and limitation section. No additional experiments have been performed to substantiate the envisioned allosteric activation mechanism of the co-receptor kinase BAK1 by the receptor EFR. I rewrote the public statement accordingly.

Reviewer #2 (Recommendations For The Authors):

Thanks for responding to my comments.

Reviewer #3 (Recommendations For The Authors):

The revised manuscript has fully addressed my previous concerns and is now suitable for publication in eLife.

Changes from the earlier version to the revised version

Abstract

Abstract

Transmembrane signaling by plant receptor kinases (RKs) has long been thought to involve reciprocal trans-phosphorylation of their intracellular kinase domains. The fact that many of these are pseudokinase domains, however, suggests that additional mechanisms must govern RK signaling activation. Non-catalytic signaling mechanisms of protein kinase domainshave been described in metazoans, but information is scarce for plants. Recently, a non-catalytic function was reported for the leucine-rich repeat (LRR)-RK subfamily XIIa member EFR (ELONGATION FACTOR TU RECEPTOR) and phosphorylation-dependent conformational changes were proposed to regulate signaling of RKs with non-RD kinase domains. Here, using EFR as a model, we describe a non-catalytic activation mechanism for LRR-RKs with non-RD kinase domains. EFR is an active kinase, but a kinase-dead variant retains the ability to enhance catalytic activity of its co-receptor kinase BAK1/SERK3 (BRASSINOSTEROID INSENSITIVE 1-ASSOCIATED KINASE 1/SOMATIC EMBRYOGENESIS RECEPTOR KINASE 3). Applying hydrogen-deuterium exchange mass spectrometry (HDX-MS) analysis and designing homology-based intragenic suppressor mutations, we provide evidence that the EFR kinase domain must adopt its active conformation in order to activate BAK1 allosterically, likely by supporting αC-helix positioning in BAK1. Our results suggest a conformational toggle model for signaling, in which BAK1 first phosphorylates EFR in the activation loop to stabilize its active conformation, allowing EFR in turn to allosterically activate BAK1.

Introduction

Plants and metazoans respond to extracellular signals through different sets of plasma membrane receptors. Whereas the G-protein coupled receptor family expanded in metazoans, the receptor kinase (RK) and receptor protein (RP) families expanded in plants (Shiu and Bleecker, 2003, 2001; Trusov and Botella, 2016). Plant RKs resemble metazoan receptor tyrosine kinases (RTKs) in their organization, containing an extracellular ligand-sensing domain that is coupled to an intracellular kinase domain by a single-pass transmembrane helix (Hohmann et al., 2017; Lemmon and Schlessinger, 2010). In contrast to RTKs, however, the intracellular kinase domain of plant RKs is a serine/threonine kinase domain most closely related to interleukin-1 receptor associated kinases (IRAKs)/Pelle kinases (Shiu and Bleecker, 2003, 2001) – although dual-specificity kinase activity has been reported in some cases (Liu et al., 2018; Luo et al., 2020; Macho et al., 2014; Oh et al., 2009; Perraki et al., 2018).

The leucine-rich repeat (LRR)-RKs are the most extensively studied sub-family of Arabidopsis thaliana (hereafter Arabidopsis) RKs. They primarily sense peptide ligands that regulate growth and development, or molecular patterns that are released during damage or microbial infection (Couto and Zipfel, 2016). Ligand-binding LRR-RKs have long ectodomains (>20 LRRs), and recruit short LRR ectodomain-containing co-receptor kinases (coRKs) upon ligand perception (Bender and Zipfel, 2023; Hohmann et al., 2017), most of which belong to the SERK (SOMATIC EMBRYOGENESIS RECEPTOR KINASE) family. One of the most well characterized LRR-RKs is BRI1 (BRASSINOSTEROID INSENSITIVE 1), which recruits the coRKs SERK1 and SERK3/BAK1 (BRI1-ASSOCIATED KINASE 1) upon brassinosteroid (BR) perception (Albrecht et al., 2008; He et al., 2000; Hothorn et al., 2011; Nam and Li, 2002; Santiago et al., 2013). After ligand perception, the intracellular kinase domains of BRI1 and BAK1 reciprocally trans-phosphorylate each other in their activation loops (Hohmann et al., 2017), and BAK1 then phosphorylates BRI1 in its juxtamembrane segment and C-tail to propagate BR signaling (Wang et al., 2008, 2005). BRI1 exemplifies a set of BAK1-dependent LRR-RKs that have RD-type intracellular kinase domains and require their catalytic activity to signal (Cao et al., 2013; Kosentka et al., 2017; Taylor et al., 2016). Other LRR-RKs such as EFR (ELONGATION FACTOR TU RECEPTOR) instead have non-RD intracellular kinases, and are thought not to trans-phosphorylate their associated coRK following ligand perception (Schwessinger et al., 2011). Moreover, kinase inactive mutants of EFR (D849N or K851E) retain signaling function, a result that challenges the generality of the reciprocal trans-phosphorylation model (Bender et al., 2021).

EFR perceives the pathogen-associated molecular pattern (PAMP) ELONGATION FACTOR TU, or its active peptide epitope elf18 (Zipfel et al., 2006). PAMP perception triggers heterodimerization with the coRK BAK1, resulting in phosphorylation of the EFR intracellular kinase domain (Bender et al., 2021; Roux et al., 2011; Schulze et al., 2010; Schwessinger et al., 2011). The signal is subsequently relayed to the cytoplasmic kinases BIK1 (BOTRYTIS-INDUCED KINASE 1) and PBL1 (PBS1-LIKE 1) (Li et al., 2014; Lu et al., 2010; Ranf et al., 2014). The resulting Immune signaling activation elicits a battery of cellular responses, including an apoplastic oxidative burst (hereafter oxidative burst), Ca2+-influx, callose deposition, MAPK (MITOGEN ACTIVATED PROTEIN KINASE) activation, and transcriptional reprogramming (DeFalco and Zipfel, 2021).

Initial studies in vitro identified phosphorylation sites in EFR that result either from auto-phosphorylation or from trans-phosphorylation by BAK1 (Wang et al., 2014). More recently, in vivo phosphorylation sites on EFR were identified by immunoprecipitating EFR-GFP from elf18-treated seedlings (Bender et al., 2021). One was a serine (S888) in the activation loop (A-loop), at which phosphorylation was consistently observed in vitro and in vivo (Bender et al., 2021; Wang et al., 2014). Despite the non-RD nature of EFR, and its ability to signal independently of catalytic activity, ligand-inducible phosphorylation in the A-loop (S887/S888) surprisingly proved indispensable for signaling (Bender et al., 2021). Furthermore, a functionally important tyrosine in EFR (Y836) – conserved in subdomain VIa of many eukaryotic protein kinases (Lai et al., 2016; Luo et al., 2020; Perraki et al., 2018) – was found to be phosphorylated after ligand treatment in vivo (Macho et al., 2014). A phospho-ablative Y836F mutation in EFR blocks immune signaling and resistance against the phytopathogenic bacterium Pseudomonas syringae (Macho et al., 2014). Thus, although EFR function does not require its catalytic activity, both A-loop phosphorylation and phosphorylation of the VIa subdomain tyrosine (VIa-Tyr) appear to be crucial. The mechanistic importance of these phosphorylation events remains poorly understood, but we hypothesized that they switch the EFR kinase domain into an active-like conformation that allosterically activates BAK1 in the EFR-BAK1 complex through a mechanism similar to that described for pseudokinases (Mace and Murphy, 2021; Sheetz and Lemmon, 2022) – ultimately promoting BAK1’s activity towards its substrate BIK1.

Here, we tested the hypothesis that EFR is an allosteric regulator of BAK1 using a range of different approaches. Our in vitro studies revealed that forcing dimerization of the EFR and BAK1 intracellular domains allosterically enhances BAK1 activity. Using a homology-guided approach, we designed mutations to stabilize the active-like conformation of the EFR kinase domain and found that they restore functionality of EFR variants that cannot be phosphorylated in the A-loop (EFRSSAA) or at Y836 (EFRY836F). We also used hydrogen-deuterium exchange mass spectrometry (HDX-MS) to analyze conformational dynamics, revealing that the Y836F mutation hampers the ability of the EFR kinase domain to adopt an active-like conformation. Collectively, our findings argue that the active conformation of the EFR kinase domain is required for allosteric activation of the BAK1 kinase domain. Finally, we present evidence suggesting that EFR activates BAK1 allosterically by supporting αC-helix positioning in BAK1.

Results

The EFR intracellular domain allosterically activates BAK1 in vitro

To test the hypothesis that EFR enhances BAK1 catalytic activity allosterically, we used rapamycin (Rap)-induced dimerization (RiD) to induce a complex between the isolated EFR and BAK1 intracellular domains (Banaszynski et al., 2005; Kim et al., 2021). The RiD system was previously applied in planta, maintaining membrane-association by N-terminal myristoylation (Kim et al., 2021). For our in vitro experiments, the myristoylation sites were excluded to facilitate the purification of recombinant protein. We tested two LRR-RKs: EFR and BRI1, fusing GFP C-terminally and FKBP (FK506 binding protein) N-terminally to their intracellular domains. BAK1 instead was N-terminally tagged with FRB (FKBP-Rap-binding). We first confirmed that adding Rap induced formation of FKBP-EFR/FRB-BAK1 dimers in size exclusion chromatography experiments (Figure 1A). We then assessed the effect of inducing RK/coRK kinase complex formation on their ability to phosphorylate BIK1D202N. BIK1 was chosen as it is a reported substrate of both, EFR/BAK1 and BRI1/BAK1 complexes (Lin et al., 2013). Rap addition increased BIK1D202N phosphorylation when the BRI1 or EFR kinase domains were dimerized with BAK1 (Figure 1B,C). Kinase-dead variants with the catalytic residue (HRD-aspartate) replaced by asparagine (EFRD849N and BRI1D1009N), had distinct effects. BRI1D1009N failed to enhance BIK1 phosphorylation substantially, whereas EFRD849N retained some ability to do so (Figure 1B,C). The same trend was observed for phosphorylation of the BAK1 kinase domain itself, indicating that EFR also enhances BAK1 autophosphorylation activity (Figure 1B,C).

EFR facilitates BIK1 trans-phosphorylation by BAK1 non-catalytically.

The kinase domains of coRK BAK1 and the RKs BRI1 and EFR were tagged with RiD domains and purified from E. coli λPP cells. A) Recombinantly expressed RiD-tagged kinase domains were mixed together at equimolar ratios (2 µM), with or without addition of 10 µM Rap as well as 10 mM MgCl2 and 1 mM AMP-PNP. B) RK and coRK were mixed at an equimolar ratio at 50 nM, kinase-dead BIK1D202N substrate was added at 500 nM. Reactions were carried out at RT for 10 min with 0.5 μCi [γ-³²P]ATP, 100 μM ATP and 2.5 mM each of MgCl2 and MnCl2. Addition of 1 μM Rap enhanced transphosphorylation of BIK1 by EFR and BRI1. Kinase-dead BRI1 failed to enhance BIK1 transphosphorylation, but kinase-dead EFR retained some ability to do so. A similar trend was observed for (auto)phosphorylation of BAK1 itself.C) Quantification of band intensities over three independent experiments of which a representative is shown in B.

The increased BIK1 trans-phosphorylation observed in these in vitro RiD experiments could arise either from direct enhancement of BAK1 activity by the RK kinase domain or from more efficient BIK1D202N recruitment to the dimerized RK/coRK complex. We found that BIK1D202N did not co-elute with the EFR-BAK1 complex in size exclusion studies, indicating that stable EFR/BAK1/BIK1 trimers do not form in vitro (Figure 1A). Our findings therefore support the hypothesis that EFR increases BIK1 phosphorylation by allosterically activating the BAK1 kinase domain. Although EFR’s catalytic activity is dispensable for this effect – and for immune signaling – EFR must be phosphorylated at S887/S888 in the A-loop and Y836 to signal in vivo (Bender et al., 2021; Macho et al., 2014). The phospho-ablative EFR A-loop mutant EFRSSAA fails to activate the RK/coRK complex, as ligand-induced phosphorylation of BAK1 S612 – a mark for active BAK1-containing receptor complexes (Perraki et al., 2018) – is obstructed despite retaining the ability to associate with BAK1 in a ligand-dependent manner (Bender et al., 2021). To test whether the requirement for Y836 phosphorylation is similar, we immunoprecipitated EFR-GFP and EFRY836F-GFP from mock- or elf18-treated seedlings and probed co-immunoprecipitated BAK1 for S612 phosphorylation. EFRY836F also obstructed the induction of BAK1 S612 phosphorylation (Figure 1 – Supplement 1), indicating that EFRY836F and EFRSSAA impair receptor complex activation.

EFR VIa-Tyr mutation affects dynamics of regulatory kinase subdomains

Y836 in EFR corresponds in sequence and structural alignments with Y156 in the canonical kinase PKA (protein kinase A) (Figure 2 – Supplement 1, 2A), which is conserved in many protein kinases. We therefore speculated that Y836 may function as a pivot for key movements of the regulatory αC-helix, controlling conformational toggling of inactive/active transitions as described for its PKA counterpart (Tsigelny et al., 1999). To test whether an EFR Y836F mutation in EFR interferes with assembly of the active-like EFR kinase conformation we used HDX-MS. Plotting differential HDX (Δ%EX) between unphosphorylated EFR and EFRY836F (aa684-1031) identifies one region that is stabilized in EFRY836F compared to wild-type and two that are destabilized (Figure 2A,B). The region stabilized in EFRY836F contains the N-terminal part of the A-loop, implying that this region makes more extensive contacts with the kinase core than in the unphosphorylated wild-type kinase domain. Alternatively, the decreased deuterium uptake in the EFRY836F A-loop could arise from stabilization of a short A-loop α-helix, as seen in the inactive conformation of many kinases (such as in epidermal growth factor receptor). The two regions that become more structurally flexible in EFRY836F include the β3-αC loop and the catalytic loop plus C-terminal end of the αE-helix – where Y836 is located (Figure 2A,B). Because these two regions are both important for conformational switching, the HDX-MS results suggest that Y836 is important in regulating kinase domain allosteric transitions. Consistent with this, crystal structures as well as AlphaFold2 models of active kinase conformations (Faezov and Dunbrack, 2023) suggest that the side-chain of this tyrosine forms hydrogen-bonds (H-bond) with the αC-β4 loop backbone to establish an inter-lobe connection (Figure 2 – Supplement 1). Loss of this inter-lobe connection may underlie the observed alterations in EFRY836F conformational dynamics in EFRY836F, resulting from inaccessibility of the active-like conformation.

EFRY836F and EFRSSAA impair the active kinase conformation, which is required for signaling function.

A) (left) HDX-MS results for unphosphorylated EFR and EFRY836F protein. The difference in percent H/D exchange in wild type EFR and EFRY836F is expressed as the Δ%EX (wild type EFR – EFRY836F), with the positive and negative Δ%EX indicating more stabilized and destabilized regions in EFRY836F, respectively, compared to wild-type EFR. The Δ%EX values at different labeling time points are shown as colored lines, as indicated in the figure. The horizontal dotted black lines indicate the 98% confidence interval for the Δ%EX data (±7.18%, corresponding to ±0.4 Da difference between wild type and Y836F percent exchange) calculated as described previously (Houde et al., 2011). Regions with Δ%EX values that exceed this confidence limit are indicated as colored bars in the figure, including the β3-αC loop (orange), the catalytic loop plus part of αE (purple), and the A-loop (blue). These regions are colored in the AlphaFold2-derived model of the EFR kinase domain shown at right, in which Y836 is shown as a purple sphere. All data are the average of three independent biological repeats (n=3) with three technical repeat experiments each. A summary of the HDX-MS analysis is presented in Table 2. B) HDX-MS analysis of representative peptides from regions with significantly different HD exchange. Frames are color-coded according to regions in A. Amino acid range of the peptides in full length EFR are indicated in the top left corner and the sequence below. C,D) Secondary site mutation EFR F761[H/M] partially restores function of EFRY836F (C) and EFRSSAA (D). Full length EFR and its variants were expressed transiently in N. benthamiana and their function was tested in an oxidative burst assay. EFR F761H partially restored oxidative bursts of EFRY836F and EFRSSAA. Outliers are in indicated by asterisk in addition to the outlier itself and are included in statistical analysis; Statistical test: Kruskal-Wallis test (p < 2.2*10-16 in C, p = 1.163*10-7 in D), Dunn’s post-hoc test with Benjamin-Hochberg correction (p <= 0.05) Groups with like lowercase letter designations are not statistically different.

Kinase activating mutations restore partial function of EFRY836F and EFRSSAA

Earlier version: To build on these results, we next sought to rescue stability of the active-like EFR kinase domain conformation by introducing activating mutations into EFRY836F. Kinase activating mutations are well known in human disease, and cause different malignancies (Foster et al., 2016; Hu et al., 2015). Because the HDX-MS data indicated a destabilized αC-helix, we were specifically interested in activating mutations thought to stabilize the kinase αC-helix in a ‘swung-in’ state, with the goal of making homologous changes in EFR. Such mutations were systematically identified in BRAF/CRAF by searching oncogenes using phenylalanine substitutions (Hu et al., 2015). We exploited these oncogenic BRAF mutations for homology-based design of putative activating EFR mutations at corresponding positions that could function as intragenic suppressors of EFRY836F (Table 1 and Figure 2 – Supplement 2A,B). Since EFR already carries a phenylalanine at the position corresponding to L505 in BRAF, we generated EFRF761[H/M] to resemble mutations known to potently activate BRAF (Hu et al., 2015).

Revised version: To build on these results, we next sought to rescue the active-like kinase conformation of EFRY836F by introducing activating mutations. Kinase activating mutations are well known in human disease, and cause different malignancies (Foster et al., 2016; Hu et al., 2015). Because the HDX-MS data indicated a destabilized αC-helix, we were specifically interested in activating mutations thought to stabilize the kinase αC-helix in a ‘swung-in’ state, with the goal of making homologous changes in EFR. Such mutations were systematically identified in BRAF/CRAF by searching oncogenes using phenylalanine substitutions (Hu et al., 2015). We exploited these oncogenic BRAF mutations for homology-based design of putative activating EFR mutations at corresponding positions that could function as intragenic suppressors of EFRY836F (Table 1 and Figure 2 – Supplement 2A,B). Since EFR already carries a phenylalanine at the position corresponding to L505 in BRAF, we generated EFRF761[H/M] to resemble mutations known to potently activate BRAF (Hu et al., 2015).

Earlier version: The putative activating mutations were introduced into either EFRWT or EFRY836F and transiently expressed in Nicotiana benthamiana to test receptor function. As expected, heterologous expression of EFRWT but not of EFRY836F conferred elf18 sensitivity (Figure 2B, Figure 2 – Supplement 2C). When introduced on their own into EFRWT, the single mutations L743F, F761[HM], and L873E were all consistent with wild-type EFR function, supporting an elf18-induced oxidative burst, except for the ΔNLLKH deletion (Figure 2 – Supplement 2C). Intriguingly, the F761[H/M] mutations were also able to partially restore the ability of EFRY836F to support an elf18-induced oxidative burst (Figure 2B and Figure 2 – Supplement 2C). This observation also extended to EFRSSAA (Figure 2C), arguing that the F761[H/M] mutations, that stabilizes the regulatory spine, support the ability of EFR to trans-activate BAK1.

Revised version: The putative activating mutations were introduced into either EFRWT or EFRY836F and transiently expressed in Nicotiana benthamiana to test receptor function. As expected, heterologous expression of EFRWT but not of EFRY836F conferred elf18 sensitivity (Figure 2B, Figure 2 – Supplement 2C). When introduced on their own into EFRWT, the single mutations L743F, F761[HM], and L873E were all consistent with wild-type EFR function, supporting an elf18-induced oxidative burst, except for the ΔNLLKH deletion (Figure 2 – Supplement 2C). Intriguingly, the F761[H/M] mutations were also able to partially restore the ability of EFRY836F to support an elf18-induced oxidative burst (Figure 2B and Figure 2 – Supplement 2C). This observation also extended to EFRSSAA (Figure 2C), arguing that the F761[H/M] mutations support the ability of EFR to trans-activate BAK1.

EFR A-loop phosphorylation was reported previously to be indispensable for EFR function, and led to the hypothesis that it controls conformational switching, despite EFR being a non-RD kinase (Bender et al., 2021). Our finding that the F761H mutation restores EFRSSAA function further supports this hypothesis. Nevertheless, how EFR A-loop phosphorylation facilitates conformational switching is unclear. In our AlphaFold2 models of the EFR kinase domain, we noted that two basic residues from the β3-αC loop and the αC-helix extend towards the A-loop and, similar to PKA H87 (Meharena et al., 2016), may coordinate A-loop phosphorylation (Figure 2 – Supplement 3).

EFR F761H restores function of EFRY836F and EFRSSAA in Arabidopsis

We next generated stable Arabidopsis complementation lines expressing pEFR::EFR(variant)-GFP::HSP18t constructs in the null efr-1 background. Because the F761H mutation showed greater ability than F761M to rescue function of both EFRY836F and EFRSSAA (Figure 2B,C), we chose this secondary mutation for the generation of complementation lines. Two independent homozygous complementation lines were isolated, with similar accumulation of EFR-GFP protein in the T3 generation. For EFRF761H/Y836F, only one homozygous T3 line could be isolated, so a heterozygous T2 line with two T-DNA insertion events was used for physiological experiments. Similar to experiments in N. benthamiana, the oxidative burst was partially restored in stable Arabidopsis complementation lines that express EFRF761H/Y836F or EFRF761H/SSAA (Figure 3 – Supplement 1A).

EFRF761H/Y836F and EFRF761H/SSAA recover receptor complex activation.

A) In infection assays, GUS activity was high in the positive control efr-1 line. GUS activity level was reduced in the EFRWT and EFRF761H complementation lines, but much less so in the EFRY836F and EFRSSAA complementation lines. By contrast, EFRF761H/Y836F and EFRF761H/SSAA complementation lines displayed substantially repressed GUS activity. Each experiment was repeated three times with similar results. Outliers are indicated by an additional asterisk and included in statistical analysis. Statistical test: Kruskal-Wallis test (p = 5.704*10-7), Dunn’s post-hoc test with Benjamin-Hochberg correction (p <= 0.05) Groups with like lowercase letter designations are not statistically different. B) In IP kinase assays, ligand-induced interaction of EFRWT and EFRF761H with BAK1 increased transphosphorylation of BIK1D202N, but this was abolished for EFRY836F and EFRSSAA. Both EFRF761H/Y836F and EFRF761H/SSAA showed partially restored BIK1D202N trans-phosphorylation as well as BAK1 S612 phosphorylation (across four replicates for EFRF761H/SSAA and in two out of four replicates for EFRF761H/Y836F). Samples were also probed for MAPK phosphorylation for effective ligand treatment. Treatment: 100 nM elf18 for 10 min. C) Quantification of BIK1D202N band intensity observed in autoradiographs from the four independent replicates performed. Dotted red line indicates unchanged band intensity in mock vs. elf18 treatment.

The oxidative burst is only one of multiple cellular responses triggered by elf18. We therefore also tested whether other immune signaling responses are restored in the EFRF761H/Y836F and EFRF761H/SSAA complementation lines. Consistent with previous results, EFRY836F and EFRSSAA lines were both less sensitive to treatment with 5 nM elf18 than those with wild-type EFR in seedling growth inhibition (SGI) assays (Bender et al., 2021; Macho et al., 2014) (Figure 3 – Supplement 1B). As expected, both independent complementation lines of EFRF761H/Y836F and EFRF761H/SSAA exhibited enhanced SGI compared to EFRY836F and EFRSSAA, respectively (Figure 3 – Supplement 1B). We also observed that MAPK activation was abolished or severely impaired in both EFRY836F and EFRSSAA complementation lines (Figure 3 – Supplement 1C) and recovered in EFRF761H/Y836F and EFRF761H/SSAA complementation lines (Figure 3 – Supplement 1C).

The recovery of multiple immune responses in EFRF761H/Y836F and EFRF761H/SSAA complementation lines suggested that they retain fully functional elf18 signaling and that effective resistance against bacteria can be established. To confirm this, we tested the transgenic lines for resistance against Agrobacterium tumefaciens. Infection by this bacterium is restricted in Arabidopsis by EFR, so loss-of-function mutants like efr-1 are more susceptible (Zipfel et al., 2006). A. tumefaciens carrying a plasmid with an intronic version of the β-glucuronidase (GUS) gene was used to evaluate infection success. GUS activity in plant protein extracts following infection correlates with the ability of the bacteria to transiently transform plant cells. Significant GUS activity was detected in efr-1 (Figure 3A), consistent with successful A. tumefaciens infection, but only little was seen in wild-type EFR complementation lines, reflecting resistance to infection (Bender et al., 2021; Zipfel et al., 2006). The EFRY836F and EFRSSAA complementation lines were both more susceptible to A. tumefaciens transformation than wild-type EFR complementation lines, as indicated by elevated GUS activity (Figure 3A), but this was greatly diminished in the EFRF761H/Y836F and EFRF761H/SSAA complementation lines. Hence, these experiments show that the EFR F761H mutation restores full signaling function of EFRY836F and EFRSSAA and thus resistance against A. tumefaciens.

EFRF761H/Y836F and EFRF761H/SSAA restore BIK1 trans-phosphorylation

As shown Figure 3 and Figure 1 – Supplement 1, EFRY836F and EFRSSAA are impaired in elf18-triggered immune signaling at the level of receptor complex activation. In both cases BAK1 S612 phosphorylation is reduced, resulting in compromised ability to trans-phosphorylate BIK1. We therefore asked next whether EFRF761H/Y836F and EFRF761H/SSAA show restored ability to induce BAK1 autophosphorylation and resulting BIK1 trans-phosphorylation. We performed semi-in vivo IP-kinase assays in which EFR-GFP variants were immunoprecipitated from two-week-old seedlings after 10 min mock or 100 nM elf18 treatment, and kinase activity of the complex was assessed by monitoring phosphorylation of recombinant His-BIK1D202N substrate. The EFRWT-GFP-BAK1 complex showed pronounced BIK1 phosphorylation compared with the unliganded control (Figure 3B, C). In contrast, although BAK1 co-precipitated with both EFRY836F-GFP and EFRSSAA-GFP after elf18 treatment, these complexes failed to increase BIK1D202N trans-phosphorylation compared with the unliganded controls. They also showed reduced levels of BAK1 phosphorylation at S612 (Figure 3B,C). However, BAK1 S612 phosphorylation (for EFRF761H/Y836F in two out of four experiments) and BIK1D202N trans-phosphorylation were partially restored for EFRF761H/Y836F and EFRF761H/SSAA in this assay (Figure 3B,C), arguing that the functional rescue observed in vivo reflects BAK1 activation effects.

Towards a mechanistic understanding of BAK1 allosteric activation by EFR

Earlier version: A phospho-ablative mutation of the BAK1 VIa-Tyr (Y403F) – analogous to Y836F in EFR – has also been reported to compromise elf18-induced signaling by interfering with receptor complex activation (Perraki et al., 2018). We recapitulated this finding using only the intracellular domains of EFR and BAK1, fusing them to FKBP and FRB domains with myristoylation sequences and inducing dimerization with Rap in N. benthamiana (Figure 4 and Figure 4 – Supplement 1A). This approach allowed us to investigate BAK1 variants without interference of endogenous NbSERKs. We wondered whether BAK1Y403F could be rescued by introducing a mutation into its regulatory spine that stabilizes the active-like confirmation, as seen for F761[H/M] in EFR. Indeed, introducing the analogous mutation (I338H) into BAK1 partly restored the oxidative burst for BAK1I338H/Y403F when dimerized with EFRWT (Figure 4 – Supplement 1A), suggesting that BAK1Y403F is perturbed, like EFRY836F, in accessing its active conformation. Analogies with allosteric regulation of other kinases and pseudokinases (Hu et al., 2013; Mace and Murphy, 2021; Sheetz and Lemmon, 2022; Zhang et al., 2006) led us hypothesize that EFR might stabilize the active-like conformation of BAK1 intermolecularly – helping to position the BAK1 αC-helix (in the EFR-BAK1 complex) to activate immune signaling. If this is correct, EFR that is partly ‘locked’ in its active-like conformation may be able to rescue (in trans) the function of a signaling-inactive BAK1 mutant with a destabilized αC-helix (as in BAK1Y403F). We therefore asked whether EFRF761H could achieve this, since it exhibited the tendency for increased BIK1D202N trans-phosphorylation in IP-kinase assays (Figure 3B,C) and for an accelerated oxidative burst in N. benthamiana (Figure 4 – Supplement 2A,B). Intriguingly, when EFRF761H was paired with BAK1Y403F, the Rap-induced oxidative burst was partially restored (Figure 4), suggesting that EFRF761H can partially restore BAK1Y403F function. The fact that no oxidative burst was seen when EFRF761H was paired with catalytically inactive BAK1D416N suggests that this restoration does not occur through a direct catalytic mechanism mediated by EFRF761H (Figure 4 – Supplement 2C), supporting the hypothesis of allosteric regulation.

Revised version: A phospho-ablative mutation of the BAK1 VIa-Tyr (Y403F) – analogous to Y836F in EFR – has also been reported to compromise elf18-induced signaling by interfering with receptor complex activation (Perraki et al., 2018). We recapitulated this finding using only the intracellular domains of EFR and BAK1, fusing them to FKBP and FRB domains with myristoylation sequences and inducing dimerization with Rap in N. benthamiana (Figure 4 and Figure 4 – Supplement 1A). This approach allowed us to investigate BAK1 variants without interference of endogenous NbSERKs. We wondered whether BAK1Y403F could be rescued by a mutation analogous to EFRF761[H/M to recover accessibility of the active conformation. Indeed, introducing the analogous mutation (I338H) into BAK1 partly restored the oxidative burst for BAK1I338H/Y403F when dimerized with EFRWT (Figure 4 – Supplement 1A), suggesting that BAK1Y403F is perturbed, like EFRY836F, in accessing its active conformation. Analogies with allosteric regulation of other kinases and pseudokinases (Hu et al., 2013; Mace and Murphy, 2021; Sheetz and Lemmon, 2022; Zhang et al., 2006) led us hypothesize that EFR might stabilize the active-like conformation of BAK1 intermolecularly – helping to position the BAK1 αC-helix (in the EFR-BAK1 complex) to activate immune signaling. If this is correct, EFR that is partly ‘locked’ in its active-like conformation may be able to rescue (in trans) the function of a signaling-inactive BAK1 mutant with a destabilized αC-helix (as in BAK1Y403F). We therefore asked whether EFRF761H could achieve this, since it exhibited the tendency for increased BIK1D202N trans-phosphorylation in IP-kinase assays (Figure 3B,C) and for an accelerated oxidative burst in N. benthamiana (Figure 4 – Supplement 2A,B). Intriguingly, when EFRF761H was paired with BAK1Y403F, the Rap-induced oxidative burst was partially restored (Figure 4), suggesting that EFRF761H can partially restore BAK1Y403F function. The fact that no oxidative burst was seen when EFRF761H was paired with catalytically inactive BAK1D416N suggests that this restoration does not occur through a direct catalytic mechanism mediated by EFRF761H (Figure 4 – Supplement 2C), supporting the hypothesis of allosteric regulation.

EFRF761H recovers BAK1Y403F function.

The cytoplasmic domains of BAK1 and EFR variants with fused RiD-tags were transiently expressed in N. benthamiana and leaf discs were treated with Rap to induce dimerization. EFR and EFRF761H induced a similar total oxidative burst when BAK1 was co-expressed. The co-expression of BAK1Y403F and EFR diminished the oxidative burst, which was restored partially when EFRF761H was co-expressed. Outliers are indicated by an additional asterisk and included in statistical analysis. Statistical test: Kruskal-Wallis test (p < 8.516 *10-7), Dunn’s post-hoc test with Benjamin-Hochberg correction (p <= 0.05) Groups with like letter designations are not statistically different.

Catalytic independence is only observed for Arabidopsis sp. LRR-RK XIIa kinase domains

EFR belongs to LRR-RK subfamily XIIa (Figure 5A), of which two other RKs are functionally described immune RKs in Arabidopsis: FLS2 and XPS1 (XANTHINE/URACIL PERMEASE SENSING 1) (Gómez-Gómez and Boller, 2000; Mott et al., 2016). This subfamily has been implicated more generally in immune signaling, which is supported by induction of immune signaling of in vivo dimerized FLS2, EFR, or XPS1-LIKE 1 (FEXL1) intracellular domains with BAK1 (Kim et al., 2021). However, the ligands are not known for most subfamily-XIIa RKs (Dufayard et al., 2017; Kim et al., 2021). Because EFR does not require its catalytic activity, we wondered whether other LRR-RK XIIa kinase domains similarly function non-catalytically, in a manner similar to pseudokinases. To test this hypothesis, we fused the EFR ectodomain to the transmembrane helix and intracellular domain of different LRR-RK XIIa members to generate elf18 responsive RKs that can dimerize with BAK1/SERKs (Rhodes et al., 2021). These chimeric proteins were expressed transiently in N. benthamiana leaves, and their immune signaling function was tested by elf18 treatment in oxidative burst assays. All chimeras induced an oxidative burst (Figure 5B), except XIIa2 (the closest FLS2-related kinase), which exhibited only a very minor response. To determine whether their signaling function required kinase activity, we next tested variants with mutations in the catalytic site (replacing the putative catalytic base HRD-aspartate with asparagine). EFRD849N induced a robust oxidative burst as expected as well as FEXL1D838N and the closely related XIIa5D839N (AT3G47570). Similarly, XPS1D856N and XIIa6D840N (AT3G47090) induced an oxidative burst independently of their catalytic activity, but the total oxidative burst was reduced compared to the catalytically active variants. Protein levels of catalytic site mutants accumulated comparably to their corresponding wild-type versions (Figure 5 – Supplement 2). In contrast, the catalytically inactive FLS2D997N showed a diminished oxidative burst (Figure 5B), suggesting that it may be mechanistically distinct. Furthermore, kinetic differences for the oxidative burst of Arabidopsis XIIa kinase domains were observed. XPS1D856N and XIIa6D840N) showed a delayed oxidative burst (Figure 5 – Supplement 1A,B). Also, EFRD849N exhibited a delay of approximately five to seven minutes compared to wild-type EFR. In contrast, oxidative bursts induced by FEXL1D838N and XIIa5D839N were only slightly delayed compared to their corresponding catalytically active kinase domains. We corroborated these findings for XIIa5 by performing in vitro kinase assays, observing that BIK1D202N and BAK1 phosphorylation increased after Rap application as effectively for XIIa5D839N as for wild-type XIIa5 (Figure 5 – Supplement 1C,D). Taken together, our results imply that catalytic activity for XIIa5 is almost fully dispensable.

Related EFR kinases from LRR-RK XIIa in the Arabidopsis genus can function independent of their calatytic activity.

A) Phylogenetic analysis of LRR-RK subfamily XIIa. Selected LRR-RK XIIa kinase domains are labeled and highlighted with purple points. The EFR-like clade contains all Arabidopsis XIIa kinases except FLS2 and XIIa2 and also selected XIIa kinases from Arabidopsis lyrata and Brassica rapa. B,C) The ectodomain of EFR was fused to the transmembrane and intracellular domain of selected LRR-RK XIIa members to create elf18-responsive chimeras for testing the immune signaling function and catalytic dependency of the related kinase domains. The chimeras were transiently expressed in N. benthamiana and tested in oxidative burst assays. All Arabidopsis LRR-RK XIIa members induced an oxidative burst except XIIa2, the closest FLS2 related kinase in the subfamily. Catalytic dependency of the kinase domains appears to vary from kinase to kinase, with catalytically dead versions of EFR, FEXL1 and XIIa5 inducing a WT-like oxidative burst and XPS1 and XIIa6 displaying a reduced oxidative burst. FLS2 kinase dead exhibited a diminished oxidative burst. Experiments were repeated three times with similar results.

Because most Arabidopsis LRR-RK XIIa kinase domains execute their signaling function at least to some extent independently of their catalytic activity, we went on to test XIIa kinase domains from other species, that we selected from a phylogenetic tree of LRR-RK XIIa kinases built from previous phylogenetic analysis (Figure 5A, Dufayard et al., 2017). The selected catalytically active XIIa kinase domains from Oryza sativa (XA21), Populus trichocarpa (PtXIIa), Solanum lycopersicum (SlXIIa), Glycine max (GmXIIa), Brassica rapa (BrEFR), and Arabidopsis lyrata (AlEFR), induced an oxidative burst (Figure 5C; Figure 5 – Supplement 1A,B). However, catalytic site mutants of all non-Arabidopsis XIIa RKs accumulated similar to their wild-type counterparts and abolished their function in oxidative burst assays (Figure 5C Figure 5 – Supplement 2).

Discussion

Our study provides insights into the non-catalytic activation mechanisms of the EFR-BAK1 complex. Rather than requiring its kinase activity to transactive BAK1, our results suggest that the EFR kinase domain undergoes a ‘conformational toggle’ to an active-like state that permits allosteric trans-activation of BAK1. EFR function is impaired by perturbation of structural elements that are important for conformational switching – specifically, mutation of the subdomain VIa Y836 or ablation of A-loop phosphorylation sites – but these effects can be reversed by intragenic suppressor mutations that stabilize the active-like conformation, notably the F761H mutation.

Earlier version: Non-catalytic functions of kinases and pseudokinases are human disease-relevant, and have emerged as important signaling regulators in metazoans (Mace and Murphy, 2021; Schmidt et al., 2021; Sheetz and Lemmon, 2022). By comparison, very little is known about pseudokinases or non-catalytic kinase functions in plants, although pseudokinases are as prevalent in plant as in metazoan kinomes (Kwon et al., 2019). Our data demonstrate that EFR and potentially other Arabidopsis XIIa kinases function as allosteric regulators of BAK1 kinase activity, implying pseudokinase-like functions. Thus, our findings establish precedence for non-catalytic mechanisms, specifically allosteric regulation, in plant RK signaling.

Revised version: Non-catalytic functions of kinases and pseudokinases are human disease-relevant, and have emerged as important signaling regulators in metazoans (Eyers and Murphy, 2013; Mace and Murphy, 2021; Schmidt et al., 2021; Sheetz and Lemmon, 2022). By comparison, very little is known about pseudokinases or non-catalytic kinase functions in plants, although pseudokinases are as prevalent in plant as in metazoan kinomes (Kwon et al., 2019). Our data demonstrate that EFR and potentially other Arabidopsis XIIa kinases function as allosteric regulators of BAK1 kinase activity, implying pseudokinase-like functions. Thus, our findings establish precedence for non-catalytic mechanisms, specifically allosteric regulation, in plant RK signaling.

Phosphorylation of EFR is crucial for signaling function and potentially supports the conformational toggle. Nevertheless, our results do not resolve how phosphorylation of the A-loop or VIa-Tyr facilitate conformational toggling for EFR, and structural studies will be required to address this question. However, our HDX-MS analysis of unphosphorylated protein suggests that EFRY836F causes the active-like conformation to become inaccessible, potentially due to a lack of an inter-lobe H-bond – consistent with earlier studies describing the corresponding VIa-Tyr as a pivot for αC-helix movements in PKA (Taylor and Kornev, 2011; Tsigelny et al., 1999). Recent analysis of the PKA hinge region demonstrate the importance of its integrity for coordinated conformational changes in the N- and C-lobe (Olivieri et al., 2023; Wu et al., 2023). Whether EFR VIa-Tyr phosphorylation contributes to these coordinated conformational changes is unclear. We attempted to directly address the role of pY836 in EFR conformational dynamics using HDX-MS but were unable to produce sufficient recombinant protein with a pTyr-analog incorporated at the Y836 site. However, phosphorylation of the VIa-Tyr presumably distorts the hinge region due to its proximity to the αC-β4-loop. Alternatively, the VIa-pTyr and C-terminal end of the αE-helix must rotate outwards (Figure 2 – Supplement 1). Rotation of the αE-helix, however, would have direct impact on the positioning of the catalytic loop and RS1. Of note, functional importance of VIa-Tyr phosphorylation was assigned based on correlation between its phosphorylation and functional impairment by phenylalanine substitution in multiple plant RKs (Liu et al., 2018; Luo et al., 2020; Macho et al., 2014; Perraki et al., 2018), making it possible that functional importance of its phosphorylation confounded a structural role. Future computational protein-structure prediction techniques that consider post-translational modifications will aid in elucidating the effect of EFR VIa-Tyr phosphorylation on EFR’s conformation. Nevertheless, mutation of VIa-Tyr may be a useful tool to break pseudokinase/non-catalytic functions by impeding conformational toggling in a way that is more effective than using conventional kinase-dead mutations.

Basal BAK1 activity in vitro is low compared to its robustly enhanced activity after Rap-induced dimerization (Figure 1B,C; Figure 5 – Supplement 1C,D), which might be further reduced in planta by negative regulators (e.g. protein phosphatases) (Segonzac et al., 2014). We propose that, by increasing the local concentration, ligand-induced association of EFR and BAK1 allows partially active BAK1 to phosphorylate the adjacent EFR A-loop. This in turn allows wild-type EFR to adopt the active-like conformation that can allosterically fully activate BAK1 to promote substrate phosphorylation. The activating mutation EFR F761H potentially circumvents the requirement of A-loop phosphorylation since it partially restored EFRSSAA function. Consistently, EFRF761H showed elevated BIK1 trans-phosphorylation in IP-kinase assays and an accelerated oxidative burst (Figures 3B, C; Figure 4 – Supplement 2A, B), suggesting constitutive assembly of the active-like conformation. However, EFRF761H still requires catalytic activity of complexed BAK1 (Figure 4-Supplement 1).

Our data demonstrate that allosteric BAK1 kinase activation plays a key role in EFR-BAK1-mediated immune signaling. The partial recovery of BAK1Y403F, which we hypothesize is impaired in αC-helix positioning, by EFRF761H suggests a mechanism for allosteric regulation involving BAK1 αC-helix positioning. Indeed, multiple metazoan kinases are allosterically regulated by αC-helix positioning (Sheetz and Lemmon, 2022). Moreover, BAK1 has a high propensity for a disordered αC-helix (Moffett et al., 2017), suggesting its correct positioning requires additional support. Different orientations of kinase-kinase dimers, in which allosteric regulation occurs at the αC-helix, were previously described (Mace and Murphy, 2021; Sheetz and Lemmon, 2022). Structural analysis will be required to resolve the exact interaction interface for allosteric activation in the EFR-BAK1 kinase dimer and to design interface-disrupting mutations for structure-function analysis but is beyond the scope of the present work.

Allosteric activation of BAK1/SERKs is partially preserved in Arabidopsis LRR-RK XIIa kinases (Figure 5). Whether these kinases function solely non-catalytically, as suggested by XIIa5D839N (Figure 5, Figure 5 – Supplement 1), or contribute also catalytically to signaling activation is unclear since oxidative bursts induced by all functional catalytic base mutants of XIIa kinases were delayed (Figure 5 – Supplement 1A,B). Catalytic base mutation may affect conformational dynamics, which we show are important for EFR function (Figure 2), confounding our interpretation. Inactive XIIa kinases with intact conformational dynamics may clarify the extent of catalytic contribution and could be obtained by development of selective inhibitors or engineering analog sensitive kinases. Alternatively, catalytic independency could be a recent innovation in Arabidopsis sp. XIIa kinases, that is taxonomically restricted due to a strong negative selection on the kinase domain (Man et al., 2023) but further study is required to test this hypothesis.

Taken together, our results add to growing evidence that non-catalytic functions of kinases are similar to bona fide kinases controlled by conformational switching (Sheetz et al., 2020; Sheetz and Lemmon, 2022) and further set precedence for discovering more non-catalytic mechanisms in plant RK signaling where pseudokinases are particularly prevalent (Kwon et al., 2019).

Limitations of the study

In our in vitro kinase assays using the RiD system, the proteins were freely diffusing as they were produced as soluble entities. Freely diffusing components, however, do not fully resemble the in planta mode of activation where the plasma membrane restricts the free movement. The colocalization of RK/coRK complexes with their substrates may be a supporting component of signaling activation, which we did not capture in our in vitro assays.

It is yet unclear whether EFR phosphorylates BAK1 efficiently on activation-relevant residues in the context of stabilized proximity. This information will ultimately help to deduce whether the catalytic activity of EFR contributes to signaling activation and whether even marginal catalytic activity of EFRD849N may be sufficient to support signaling activation.

Further, our study relies heavily on mutations. Mutations can have pleiotropic effects as they impact conformational plasticity which supports catalytic activity and non-catalytic functions. Thus, it is difficult to untangle catalytic from non-catalytic functions using solely mutations. Eventually, it would be beneficial to obtain structural information of native proteins in a complex and monitor site-specific and time-resolved phosphorylation on complex components.

We do not have experimental support for EFRF761H stabilizing or facilitating the transition into an active-like conformation and relieving the negative effect of the Y836F mutation in EFRF761H/Y836F. It would have been beneficial to obtain HDX-MS data for these mutants as well, but we could not produce sufficient protein to do so. Thus, we rather derived the biochemical and structural effects of the EFRF761H mutant by homology to mammalian kinases, e.g. BRAF.